<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
      <text>top/Default for Sample Questions</text>
    </category>
    <info format="moodle_auto_format">
      <text>The default category for questions shared in context 'Sample Questions'.</text>
    </info>
    <idnumber></idnumber>
  </question>


<!-- question: 55198  -->
  <question type="stack">
    <name>
      <text>Compute norm</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Compute the norms of the vectors \(u=({@u_1@},{@u_2@},{@u_3@})\) and \(v=({@v_1@},{@v_2@},{@v_3@},{@v_4@},{@v_5@})^{T}\).</p>
<p>Give your answer in exact form.</p>
<p>Hint: The square root of a number \(x\) can be typed as <b>sqrt(x)</b>.</p>
<p>\(\bf{||u||}=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>\(\bf{||v||}=\)[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked Solution</b></p>
<p>The norm of a vector is simply the square root of the sum of each component squared.</p>
<p>
\(
\begin{align} 
\left \| \bf{u} \right \| &= \sqrt{{@u_1@}^{2}+({@u_2@})^{2}+{@u_3@}^{2}} \\
                                  &={@ta_1@}
\end{align}
\)
</p>
<p>
\(
\begin{align} 
\left \| \bf{v} \right \| &= \sqrt{{@v_1@}^{2}+({@v_2@})^{2}+{@v_3@}^{2}+{@v_4@}^{2}+{@v_5@}^{2}} \\
                                  &={@ta_2@}
\end{align}
\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>4.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*Question variables*/
u_1:rand(3)-1; u_2:rand(3)+1; u_3:rand(2)+3; v_1:rand(3)+2; v_2:rand(3)-1; v_3:rand(3); v_4:rand(4)+4; v_5:rand(3)+2;
vector_u:[u_1,u_2,u_3];
vector_v:[v_1,v_2,v_3,v_4,v_5];
/*Model answers*/
ta_1:sqrt(u_1^2+u_2^2+u_3^2); /*norm u*/
ta_2:sqrt(v_1^2+v_2^2+v_3^2+v_4^2+v_5^2);  /*norm v*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text><![CDATA[<p>Compute the norms of the vectors \(\bf{u}=({@u_1@},{@u_2@},{@u_3@})\) and \(\bf{v}=({@v_1@},{@v_2@},{@v_3@},{@v_4@},{@v_5@})^{T}\).</p>

<p>\(\bf{u}={@ta_1@}\)</p>
<p>\(\bf{v}={@ta_2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*possible mistakes*/
/*Node 1 checks for correct answer*/
/*Node 2 mistake:Forgetting to find the square root*/
mistake_2:(ta_1)^2;
/*Node 3 mistake:Forgetting to square the elements of the vector.*/
mistake_3:sqrt(u_1+u_2+u_3);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have forgotten to find the square root after your computations. Note that: the norm of a vector is simply the square root of the sum of each component squared.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have forgotten to square the elements of the vectors. Note that: the norm of a vector is simply the square root of the sum of each component squared.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2 mistake:Forgetting to find the square root*/
mistake_2:(ta_2)^2;
/*Node 3 mistake:Forgetting to square the elements of the vector.*/
mistake_3:sqrt(v_1+v_2+v_3+v_4+v_5);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have forgotten to find the square root after your computations. Note that: the norm of a vector is simply the square root of the sum of each component squared.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have forgotten to square the elements of the vectors. Note that: the norm of a vector is simply the square root of the sum of each component squared.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1785040430</deployedseed>
    <deployedseed>855011139</deployedseed>
    <deployedseed>86750005</deployedseed>
    <deployedseed>550699635</deployedseed>
    <deployedseed>403152655</deployedseed>
    <deployedseed>1111611645</deployedseed>
    <deployedseed>1752208510</deployedseed>
    <deployedseed>54160178</deployedseed>
    <deployedseed>1664894298</deployedseed>
    <deployedseed>1922766105</deployedseed>
    <deployedseed>295149740</deployedseed>
    <deployedseed>1291772802</deployedseed>
  </question>


<!-- question: 54786  -->
  <question type="stack">
    <name>
      <text>Simple equation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>1. Find the solution of \(2x-3=9-4x\)</p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>1.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text>x</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 55200  -->
  <question type="stack">
    <name>
      <text>Compute inner product</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given two vectors  \(u\) and \(v\) such that</p>
<p>\(u=\bf {@u@}\) and \(v=\bf {@v@}\) </p>
<p>Let <\(u, v\)> denote the inner product.</p>
<p>Compute: </p>
<p><\(u, v\)>\(=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p>First, we will get the transpose of \(u\) by interchanging its columns into rows.</p>
<p>\(u^T={@u_transpose@}\)</p>
<p>Thereafter we will multiply the transpose of vector \(u\) and the vector \(v\) as shown below inorder to get the inner product.</p>
<p>\(u^T\cdot v={@u_transpose@}\cdot {@v@}\)</p>
<p>\(\quad\quad={@u_transpose[1,1].v[1,1]@}+{@u_transpose[1,2].v[2,1]@}+{@u_transpose[1,3].v[3,1]@}\)</p>
<p>\(\quad\quad={@inner_pr@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question variables*/
u:rand(matrix([3],[5],[2]));
v:rand(matrix([2],[3],[4]));
u_transpose:transpose(u);
v_transpose:transpose(v);
aa:u_transpose[1,1].v[1,1];
bb:u_transpose[1,2].v[2,1];
cc:u_transpose[1,3].v[3,1];
dd:aa+bb+cc;
/*model answer*/
inner_pr:u_transpose.v;/*inner product*/
outer_pr:u.v_transpose;/*outer product*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text>{@inner_pr@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>inner_pr</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>inner_pr</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1963668641</deployedseed>
    <deployedseed>1492114889</deployedseed>
    <deployedseed>98172293</deployedseed>
    <deployedseed>261939880</deployedseed>
    <deployedseed>1478311376</deployedseed>
    <deployedseed>1929891303</deployedseed>
    <deployedseed>455255506</deployedseed>
    <deployedseed>1656145340</deployedseed>
    <deployedseed>1767063591</deployedseed>
    <deployedseed>229523713</deployedseed>
    <deployedseed>4198645</deployedseed>
    <deployedseed>1167808177</deployedseed>
  </question>

</quiz>